Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```scala
val NAME = ""
val COLLABORATORS = ""
```

---

# Lab2 Sequential Logic
> Labs will be due each week before the homeworks. They are not intended take a significant amount of time but rather to provide examples/practice on specific and isolated features in the language. Labs are autograded so you can get quick feedback.

### Import the necessary Chisel dependencies. 
> There will be a cell like this in every lab. Make sure you run it before proceeding to bring the Chisel Library into the Jupyter Notebook scope!

```scala
val path = System.getProperty("user.dir") + "/resource/chisel_deps.sc"
interp.load.module(ammonite.ops.Path(java.nio.file.FileSystems.getDefault().getPath(path)))
```

```scala
import chisel3._
import chisel3.util._
import chiseltest._
import chisel3.tester.RawTester.test
```

## Problem 1 (2 pts) - State
> We can use registers to store state across multiple cycles. In Chisel, to build registers we use `Reg` as well as `RegNext` and `RegInit` ([more info](https://www.chisel-lang.org/chisel3/docs/explanations/sequential-circuits.html)). Fill in the `Delay2` module such that `out` signal is equal to the `in` signal delayed by 2 cycles.

<!-- grade_id: delay2, points: 0, grade: false, solution: true -->
```scala
class Delay2 extends Module {
    val io = IO(new Bundle {
        val in  = Input(UInt(5.W))
        val out = Output(UInt(5.W))
    })
    
    // YOUR CODE HERE
    ???
}
```

<!-- grade_id: test-delay2, points: 2, grade: true, solution: false -->
```scala
def testDelay2: Boolean = {
    test(new Delay2) { dut =>
        // Cycle0
        dut.io.in.poke(5.U)
        dut.io.out.expect(0.U)
        dut.clock.step(1)
        
        
        // Cycle1
        dut.io.in.poke(4.U)
        dut.io.out.expect(0.U)
        dut.clock.step(1)
        
        
        // Cycle2
        dut.io.in.poke(3.U)
        dut.io.out.expect(5.U)
        dut.clock.step(1)
        
        // Cycle3
        dut.io.out.expect(4.U)
        dut.clock.step(1)
        
        // Cycle4
        dut.io.out.expect(3.U)
        dut.clock.step(1)
    }
    true
}
assert(testDelay2)
```

## Problem 2 (3 pts) - Accumulator
> Let's build an _accumulator_. Each cycle `en` is high, it will add `in` to it's internal total. The internal total is visible as the output `out`. The internal total should initialize to 0 on reset.

<!-- grade_id: Accumulator, points: 0, grade: false, solution: true -->
```scala
class Accumulator(n: Int) extends Module {
    val io = IO(new Bundle {
        val in  = Input(UInt(n.W))
        val en  = Input(Bool())
        val out = Output(UInt(n.W))
    })
    
    // YOUR CODE HERE
    ???
    
}
```

<!-- grade_id: test-accumulator, points: 3, grade: true, solution: false -->
```scala
def testAccumulator: Boolean = {
    test(new Accumulator(4)) { dut =>
        // Cycle 0
        dut.io.in.poke(0.U)
        dut.io.en.poke(0.B)
        dut.clock.step(1)
        dut.io.out.expect(0.U)

        // Cycle 1
        dut.io.in.poke(0.U)
        dut.io.en.poke(0.B)
        dut.clock.step(1)
        dut.io.out.expect(0.U)

        // Cycle 2
        dut.io.in.poke(3.U)
        dut.io.en.poke(0.B)
        dut.clock.step(1)
        dut.io.out.expect(0.U)

        // Cycle 3
        dut.io.in.poke(3.U)
        dut.io.en.poke(1.B)
        dut.clock.step(1)
        dut.io.out.expect(3.U)

        // Cycle 3
        dut.io.in.poke(4.U)
        dut.io.en.poke(1.B)
        dut.clock.step(1)
        dut.io.out.expect(7.U)        
    }
    true
}

assert(testAccumulator)
```

## Problem 3 (2 pts) - Scala Loops
> To familiarize ourselves with looping in Scala, we will attempt a variant of the game [_Fizz buzz_](https://en.wikipedia.org/wiki/Fizz_buzz):

* Count from 1 until `max` (exclusive)
* Maintain three sums: `fizz`, `buzz`, and `fizzbuzz`
* If the current count is divisible by ...
    * `15` add count to `fizzbuzz`
    * `5` add count to `buzz`
    * `3` add count to `fizz`
    * Add to the sums with the precedence rules given, so for example, if the count is `15`, only add to `fizzbuzz`.

<!-- grade_id: fizzbuzz, points: 0, grade: false, solution: true -->
```scala
def fizzbuzz(max: Int): (Int, Int, Int) = {
    var (fizz, buzz, fizzbuzz) = (0, 0, 0)
    
    // YOUR CODE HERE
    ???
}

fizzbuzz(20)
```

<!-- grade_id: testfizzbuzz, points: 2, grade: true, solution: false -->
```scala
assert(fizzbuzz(2) == (0, 0, 0))
assert(fizzbuzz(20) == (48, 15, 15))
assert(fizzbuzz(35) == (153, 60, 45))

```

## Problem 4 (3 pts) - ROM
> Vecs allow us to dynamically index into a Chisel collection during operation in the generated hardware. The `ROM` module provided below stores 32 read-only UInts (0.U to 31.U) and selects one to output via the `select` signal. Fill in the tester `testROM` to exhaustively test all of `ROM`'s inputs.

```scala
class ROM extends Module {
    val io = IO(new Bundle {
        val select = Input(UInt(5.W))
        val out = Output(UInt(5.W))
    })
    
    // 0.U, 1.U, ..., 31.U
    val romData = Seq.tabulate(32)(i => i.U(5.W))
    show(s"romData: $romData")
    
    val ROM: Vec[UInt] = VecInit(romData)
    
    io.out := ROM(io.select)
}
```

<!-- grade_id: testROM, points: 0, grade: false, solution: true -->
```scala
def testROM: Boolean = {
    test(new ROM) { dut =>
        
        // YOUR CODE HERE
        ???
    }
    true
}
```

<!-- grade_id: test-testROM, points: 3, grade: true, solution: false -->
```scala
assert(testROM)
```

## Problem 5 (5 pts) - Parameterized ROM
> In this problem we will revise the module `ROM` from the previous problem, so ensure your tester for it is complete and correct first. Write a `ROM2` module below that takes a Scala Int `size` as a parameter. `ROM2` will generates `size` read-only UInts stored as a `Vec`. The `select` input signal will determine which UInt to set output `out` to. Ensure the widths are correct. `ROM2` behaves the same as `ROM`, except its size is a configurable parameter. You will also need to make a tester for it (`testROM2`).

<!-- grade_id: Rom2, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
???
```

<!-- grade_id: testRom2, points: 0, grade: false, solution: true -->
```scala
def testROM2: Boolean = {
    for (size <- Seq(16, 17, 32, 33)) {
        test(new ROM2(size)) { dut =>

            // YOUR CODE HERE
            ???
        }
    }
    true
}
```

<!-- grade_id: test-testRom2, points: 5, grade: true, solution: false -->
```scala
assert(testROM2)

```

## Problem 6 (2 pts) - Scala functions
> We can write Chisel logic within a Scala function and use them within our modules. Write a Scala function named `sum` that takes two UInts as arguments and returns the sum (with width growth). You might find the [Chisel Cheat Sheet](https://github.com/freechipsproject/chisel-cheatsheet/releases/latest/download/chisel_cheatsheet.pdf) a helpful reference.

<!-- grade_id: ScalaCondPractice, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
???
```

<!-- grade_id: testSum, points: 2, grade: true, solution: false -->
```scala
def testSum(): Boolean = {
    test(new Module {
            val io = IO(new Bundle {
                val in1 = Input(UInt(4.W))
                val in2 = Input(UInt(4.W))
                val out = Output(UInt())
            })
            io.out := sum(io.in1, io.in2)
        }) { dut => 
            for (in1 <- 0 until 16) {
                for (in2 <- 0 until 16) {
                    dut.io.in1.poke(in1.U)
                    dut.io.in2.poke(in2.U)
                    dut.io.out.expect((in1 + in2).U)
                }
            }
    }
    true
}

assert(testSum)
```

## Problem 7 (5 pts) - Histogram
> Let's put `Reg`s and `Vec`s together to build a `Histogram` generator. The generated hardware will maintain `n` (a parameter) sums. Inside the generated hardware, there will be `n` registers. Each cycle the `sel` input chooses which register to add `in` input to. The `out` output should provide the new accumulation value that will written into the register by next cycle. In other words, `out` has no delay to show the new value.

<!-- grade_id: Histogram, points: 0, grade: false, solution: true -->
```scala
class Histogram(n: Int) extends Module {
    val io = IO(new Bundle {
        val in  = Input(UInt(1.W))
        val sel = Input(UInt(log2Ceil(n).W))
        val out = Output(UInt(5.W))
    })
    
    // YOUR CODE HERE
    ???
    
}
```

<!-- grade_id: test-histogram, points: 5, grade: true, solution: false -->
```scala
def testHistogram: Boolean = {
    test(new Histogram(2)) { dut =>
        // Cycle0
        dut.io.in.poke(1.U)
        dut.io.sel.poke(0.U)
        dut.io.out.expect(1.U)
        dut.clock.step(1)
        
        
        // Cycle1
        dut.io.in.poke(1.U)
        dut.io.sel.poke(0.U)
        dut.io.out.expect(2.U)
        dut.clock.step(1)
        
        
        dut.io.in.poke(0.U)
        dut.io.sel.poke(1.U)
        dut.io.out.expect(0.U)
        dut.clock.step(1)
        
        // Cycle 3
        dut.io.in.poke(1.U)
        dut.io.sel.poke(1.U)
        dut.io.out.expect(1.U)
    }
    true
}

assert(testHistogram)
```