Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```scala
val NAME = ""
val COLLABORATORS = ""
```

---

# Lab 5 - Inheritence
> Labs will be due each week before the homeworks. They are not intended take a significant amount of time but rather to provide examples/practice on specific and isolated features in the language. Labs are autograded so you can get quick feedback.

### Import the necessary Chisel dependencies. 
> There will be a cell like this in every lab. Make sure you run it before proceeding to bring the Chisel Library into the Jupyter Notebook scope!

```scala
val path = System.getProperty("user.dir") + "/resource/chisel_deps.sc"
interp.load.module(ammonite.ops.Path(java.nio.file.FileSystems.getDefault().getPath(path)))
```

```scala
import chisel3._
import chisel3.util._
import chiseltest._
import chisel3.tester.RawTester.test
```

## Problem 1 (5 pts) - Method Templated for Chisel Types
> Write a generic `connect` method that takes in two generic arguments `a` and `b` of type `chisel3.Data` and connects `a` (input) to `b` (output).

<!-- grade_id: prob2, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
???
```

<!-- grade_id: prob2-test, points: 5, grade: true, solution: false -->
```scala
class TestConnect extends Module {
    val io = IO(new Bundle {
        val u = Output(UInt())
        val s = Output(SInt())
    })
     
    connect(io.u, 10.U)
    connect(io.s, -10.S)
}

def testConnect: Boolean = {
    test(new TestConnect) { dut =>
        dut.io.u.expect(10.U)
        dut.io.s.expect(-10.S)
    }
    true
}
assert(testConnect)
```

## Problem 2 (7 pts) - Case Class Pattern Matching and Traits
> Fill in the missing `eatFry`, `takeDrink`, and `eatItem` methods in the following traits using only pattern matching (e.g. `match` with `case` and no if/else chains. Fill in the missing body of `Meal` using pattern matching such that it attempts to populate the fields added by the traits based on `p` enabling the `eatMeal` method to behave as specified by the test cases. If someone tries to add an entree not on the Menu, it should set the entry to `"none"`.

<!-- grade_id: prob1, points: 0, grade: false, solution: true -->
```scala
case class MealChoices(drink: Boolean, fries: Boolean, entree: String)

trait Fries {
    val hasFries: Boolean
    
    // YOUR CODE HERE
    ???
}

trait Drink {
    val hasDrink: Boolean
    
    // YOUR CODE HERE
    ???
}

trait Entree {
    val item: String
    
    // YOUR CODE HERE
    ???
}

object Menu {
    val validEntrees = Seq("Cheeseburger", "Quesadilla", "Hard Taco")
}

class Meal(p: MealChoices) extends Entree with Fries with Drink {
    // YOUR CODE HERE
    ???

    def eatMeal: String = s"$eatFry, $takeDrink, $eatItem"
}
```

<!-- grade_id: prob1-test, points: 7, grade: true, solution: false -->
```scala
val m0 = new Meal(MealChoices(drink=true, fries=false, entree="Hard Taco"))
assert (m0.eatMeal == "no fries :(, ahh, yummy Hard Taco", "fail m0")

val m1 = new Meal(MealChoices(drink=false, fries=true, entree="Hard Taco"))
assert (m1.eatMeal == "mmm, so thirsty :(, yummy Hard Taco", "fail m1")

val m2 = new Meal(MealChoices(drink=false, fries=false, entree="Hard Tac0"))
assert (m2.eatMeal == "no fries :(, so thirsty :(, so hungry :(", "fail m2")

val m3 = new Meal(MealChoices(drink=true, fries=true, entree="Cheeseburger"))
assert (m3.eatMeal == "mmm, ahh, yummy Cheeseburger", "fail m3")
```

## Problem 3 (10 pts) - Chisel trait
> Fill in the missing fields in the `DbgCounter` trait, so its usage below will pass its corresponding test. The purpose of this trait is to add a debug counter that prints out the current cycle if the `debug` input (Bool) is true. The trait should add an output `count` (UInt) representing the debug counter.

<!-- grade_id: prob6, points: 0, grade: false, solution: true -->
```scala
trait DbgCounter extends MultiIOModule {
    def n: Int

    // YOUR CODE HERE
    ???
}


```

<!-- grade_id: prob6-tests, points: 10, grade: true, solution: false -->
```scala
class TestMod(maxCycles: Int) extends MultiIOModule with DbgCounter {
   def n = maxCycles 
}

def testDbgCounter(n: Int, printDbg: Boolean): Boolean = {
    test(new TestMod(n)) { dut =>
        dut.debug.poke(printDbg.B)
        for (i <- 0 until n) { 
            dut.count.expect(i.U)
            dut.clock.step()
        }
        // test that counter wraps
        dut.count.expect(0.U)
    }
    true
}

assert(testDbgCounter(n=5, true))
assert(testDbgCounter(n=15, false))
```